# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['one', 'two', 'three', 'four'])
    2
    """
    
    longest_pos = 0
    longest_len = 0
    
    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):
        
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))
            
    return longest_pos

### END SOLUTION
```

```python
def longest(my_list):
    """(list)->int
    Find the longest item in a list of strings and return the index for that string.
    If there are multiple strings with the same length, return the index of the first of those longest strings.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['abc','def','ghi'])
    0
    """

    if type(my_list) != list:
        raise TypeError("That isn't a list!")
    
    if len(my_list) == 0:
        raise ValueError("There were no items in that list!")
    
    max_length = -1
    max_index = None

    for num, item in enumerate(my_list):
        try:
            if len(item) > max_length:
                max_length = len(item)
                max_index = num
        except:
            pass
        
    return max_index
    
```

```python
longest(['happy', 'birthday', 2, 'me'])
```

Output:
```
1
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['abc','def','ghi'])
Expecting:
    0
ok
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c (temp_f):
    celsius = (temp_f - 32)*5/9
    print("local variable celsius ="+ str(celsius))
    return celsius

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

Output:
```
local variable celsius =100.0
local variable celsius =0.0
local variable celsius =37.0
```

```python
f_to_c(212)
```

Output:
```
local variable celsius =100.0
100.0
```

```python
f_to_c(32)
```

Output:
```
local variable celsius =0.0
0.0
```

```python
f_to_c(98.6)
```

Output:
```
local variable celsius =37.0
37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

```python
f_to_c(32)
```

Output:
```
local variable celsius =0.0
0.0
```

```python
f_to_c(98.6)
```

Output:
```
local variable celsius =37.0
37.0
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python

```

```python
def los(admit,discharge,diagnosis):
    """(date,date,str) -> int
    This function will calculate the length of stay in days for each encounter
    If the an encounter has a diagnosis of 'observation'regardless of number
    
    >>>los(date(2019,1,1), date(2019,1,11), "CHF")
    10
    
    >>> los(date(2022,5,10), date(2022,5,30), "observation")
    0
    """
    if diagnosis == "observation":
        return 0
    else:
        return (discharge-admit).days
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
/opt/tljh/user/lib/python3.6/doctest.py in _parse_example(self, m, name, lineno)
    701         # indented; and then strip their indentation & prompts.
    702         source_lines = m.group('source').split('\n')
--> 703         self._check_prompt_blank(source_lines, indent, name, lineno)
    704         self._check_prefix(source_lines[1:], ' '*indent + '.', name, lineno)
    705         source = '\n'.join([sl[indent+4:] for sl in source_lines])

/opt/tljh/user/lib/python3.6/doctest.py in _check_prompt_blank(self, lines, indent, name, lineno)
    788                                  'lacks blank after %s: %r' %
    789                                  (lineno+i+1, name,
--> 790                                   line[indent:indent+3], line))
    791 
    792     def _check_prefix(self, lines, prefix, name, lineno):

ValueError: line 5 of the docstring for NoName lacks blank after >>>: '    >>>los(date(2019,1,1), date(2019,1,11), "CHF")'
```

```python
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for e, v in encounters.items():
    print(los(v['admit'],v['discharge'],v['diagnosis']))
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

### BEGIN SOLUTION
def average_los(test):
    """If the length of stay of any individual encounter is 0, then it will not be counted toward the average.
    >>> average_los(test1)
    5
    """
    number = 0
    encounters = len(encounters)
    average_los = 0
    for enc, los in test.item():
         average_los = (los["discharge"] - los["admit"]).days
    if average_los > 0:
        number = number + average_los
        encounters = encounter + 1
        average = number/encounter
    return average


### END SOLUTION
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f (temp_c):
    """(float>> float)
     This function converts the temperature provided in degree celsius
    
        The function is run taking the below provided data values
        
     >>> c_to_f(100)
        212
    
    
     >>> c_to_f(0)
        32.0
    
     >>> c_to_f(37)
       98.6
       """
    
    Fahrenheit = (temp_c * 9/5) + 32
    
    return Fahrenheit

```

```python
c_to_f(100)
```

```python
c_to_f(0)
```

```python
c_to_f(37)
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

# import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)


```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git add week05_assignment_2.ipynb
    !git commit -a -m "Submitting the week 5 programming exercises"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```


---

If the message above says something like _Submitting the week 3 review exercises_ or _Everything is up to date_, then your work was submitted correctly.

```python

```